<font size="4" style="color:red;"> **IMPORTANT: ** When submitting this homework notebook, please modify only the cells that start with:</font>

```python
# modify this cell
```

# Conditional Probability and Baye's Rule


## Problem

There are two urns $A$ and $B$. Urn $A$ contains $r_A$ red balls and $w_A$ white balls whereas urn $B$ contains $r_B$ red balls and $w_B$ white balls. One of the urns is picked at random and then one ball is picked at random from this urn. Write a function **conditional_probability** that calculates the conditional probability that the randomly chosen ball belonged to urn $A$ given that it is white. Assume that $\frac{r_A}{w_A}\neq\frac{r_B}{w_B}$.

<font  style="color:blue"> **Code**</font>
```python
rA, wA, rB, wB = 1., 2., 2., 1.
conditional__probability(rA, wA, rB, wB) 
```

<font  style="color:magenta"> **Output**</font>
```
0.6666666666666666
```

```python
# modify this cell

def conditional__probability(rA, wA, rB, wB):
    # inputs: all of them are of type 'float'
    # output: a variable of type 'float'
    
    P= wA/(2*(rA+wA))
    return(P)

```

<!-- grade_id: ex1, points: 5, grade: true, solution: false -->
```python
print(conditional__probability(2., 4., 3., 3.))

assert( abs(conditional__probability(2., 4., 3., 3.) -0.5714285714285715) < 10**-5)  
assert( abs(conditional__probability(1., 3., 5., 2.) -0.7241379310344829) < 10**-5) 

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

Output:
```
0.3333333333333333
AssertionError                            Traceback (most recent call last)
<ipython-input-14-0d4f3fe6a908> in <module>()
      1 print(conditional__probability(2., 4., 3., 3.))
      2 
----> 3 assert( abs(conditional__probability(2., 4., 3., 3.) -0.5714285714285715) < 10**-5)
      4 assert( abs(conditional__probability(1., 3., 5., 2.) -0.7241379310344829) < 10**-5)
      5 

AssertionError:
```

```python








```